<div class="alert alert-block alert-info">
<b>Number of points for this notebook:</b> 0.2
<br>
<b>Deadline:</b> October 21, 2025 (Tuesday) 23:59 (end of day)
</div>


# Rules of the course and other important information

In order to participate in the course, you have to accept the rules of the course.

### Important information about the course

* Please always follow carefully the schedule of the course and the deadlines on ILIAS.
* Please review the grading policy. Note that we have a zero-point policy for late submissions.
* Please keep in mind that **the most important learning outcome of these exercises is that you can debug your deep learning code by yourself**, which means that you should learn how to test your implementation. The notebooks contain lots of examples of unit tests and we hope that you will learn how to design tests during this course.

### Discussion etiquette

* You are encouraged to participate in discussions, help others learn pytorch, give tips which may help others fix problems etc. However, please be considerate and remember that the assignments are **individual projects**, not a group work.

* We kindly ask you:
  * Not to share pieces of working solution (code).
  * Not to share links to public repositories from which part of the code can be borrowed.
  * Not to share the math needed to solve some assignments. Doing the math is part of the assignment.
  * Not to share tests that can validate the correctness of a solution.

* If someone reports a problem for which you know the solution, you are welcome to give a tip but please do not give a direct answer explaining how to fix the problem. For example, instead of saying "*I had exactly the same problem and the reason was that my collate function did not do proper padding in the target sequences*", you could say "*The problem might be in your collate function, you should probably test that*". Please remember that a student learns if s/he find the problem her/himself.

* **Please note that the teaching staff cannot provide personal assistance, such as looking at your code, outside exercise sessions.**
* If you want a teacher to look at your code, please attend the exercise sessions.

### Terms and conditions

* You give permission for proctoring your submissions.
* Solution sharing is strictly not allowed **before, during and after the course**. That means that you are not allowed to share your solutions **(or any parts of those)** via private channels and/or public repositories.

```python
I_accept_the_rules = False  # Set this flag to True to accept the rules
```

<!-- grade_id: cell-3b12a7cdb8ebd4f1, points: 0.2, grade: true, solution: false -->
```python
assert I_accept_the_rules == True, 'In order to participate in the course, you have to accept the rules of the course'
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [2], in <cell line: 1>()
----> 1 assert I_accept_the_rules == True, 'In order to participate in the course, you have to accept the rules of the course'

AssertionError: In order to participate in the course, you have to accept the rules of the course
```